# DSP Lab 2022-23


This lab assignment requires preparation before the lab session.

Paragraphs that start with ❓explain what you will need to do in the following cell.

Before you turn this notebook in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).


---

## Proj2 - Polyphonic Synthesis

In this second part of the project, it is intended to use a sequence of musical notes, duration, and intensity to generate a synthesized version of the same melody in a similar way to what was done in the Lab2 assignment. In that work, a monophonic synthesizer was used which can only produce one note at a time. A polyphonic synthesizer can start playing a different note before the previous one ends, and it can also produce chords. The ability to start a note before the end of the previous one is essential for creating sounds similar to those produced by a piano or a guitar. The [PySynth](https://mdoege.github.io/PySynth/) system page presents several examples of audio signals produced by overlaying simple synthesis modules.

### Group and Student ID

❓Initialize the variable `group_id` with the number that Fenix assigned to your group and `ist_id1` and `ist_id2` with your student numbers.

<!-- grade_id: cell-2acee589c7979f8e, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()

print(f"Group number: {group_id}")
print(f"Student number: {ist_id1}")
print(f"Student number: {ist_id2}")
```

<!-- grade_id: cell-af624e6642d3e7fd, points: 10, grade: true, solution: false -->
```python
assert isinstance(group_id, int) and isinstance(ist_id1, int) and isinstance(ist_id2, int)
assert (group_id > 0) and (group_id < 35)
assert (ist_id1 > 60000) and (ist_id1 < 120000) and (ist_id2 > 60000) and (ist_id2 < 120000)
```

### Instructions and Description

❓Explain how to use your code and describe how it works.

<!-- grade_id: cell-8cfedd49f305e6c7, points: 50, grade: true, solution: true -->
YOUR ANSWER HERE

### Code

❓Place you code in the next cells. You can include additional code and text cells.



```python

```

```python

```

```python

```

### Assessment

❓Assess the quality of the synthesized audio files. Explain the current limitations of the system and what can be done to improve it.

<!-- grade_id: cell-13b1cef5f7407fdf, points: 50, grade: true, solution: true -->
YOUR ANSWER HERE

### Code tests

❓Use the next cells to include tests of your code. Call the functions that you have developed with known inputs and compare its return values with the expected results.



```python

```

```python

```