# DSP Lab 2022-23


This lab assignment requires preparation before the lab session.

Paragraphs that start with ❓explain what you will need to do in the following cell.

Before you turn this notebook in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).


---

## Proj1 - Onset Detection

This first part of the project is dedicated to the automatic location of the beginning of each note in recordings of melodies played by a musical instrument or sung by a human voice. The system can be optimized for a particular instrument or singing style. The melody to be analyzed must have at least 15 musical notes and the system must have a reasonable performance in, at least, 3 different recordings. The chosen recordings must have an annotation of the onsets that can be annotated by the students using, for example, the audio editing software [Audacity](https://www.audacityteam.org). The article [Faghih et al. (2022)](https://www.mdpi.com/2076-3417/12/15/7391) describes several tools for the automatic detection of onsets and performance assessment.

#### References
Faghih, Behnam, Sutirtha Chakraborty, Azeema Yaseen, and Joseph Timoney. ‘A New Method for Detecting Onset and Offset for Singing in Real-Time and Offline Environments’. _Applied Sciences_ 12, no. 15 (January 2022): 7391. [https://doi.org/10.3390/app12157391](https://doi.org/10.3390/app12157391).

### Group and Student ID

❓Initialize the variable `group_id` with the number that Fenix assigned to your group and `ist_id1` and `ist_id2` with your student numbers.

<!-- grade_id: cell-2acee589c7979f8e, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()

print(f"Group number: {group_id}")
print(f"Student number: {ist_id1}")
print(f"Student number: {ist_id2}")
```

<!-- grade_id: cell-af624e6642d3e7fd, points: 10, grade: true, solution: false -->
```python
assert isinstance(group_id, int) and isinstance(ist_id1, int) and isinstance(ist_id2, int)
assert (group_id > 0) and (group_id < 35)
assert (ist_id1 > 60000) and (ist_id1 < 120000) and (ist_id2 > 60000) and (ist_id2 < 120000)
```

### Instructions and Description

❓Explain how to use your code and describe how it works.

<!-- grade_id: cell-8cfedd49f305e6c7, points: 50, grade: true, solution: true -->
YOUR ANSWER HERE

### Code

❓Place you code in the next cells. You can include additional code and text cells.



```python

```

```python

```

```python

```

### Assessment

❓Present the perfomance of the onset detector on the test signals and comment the achieved results

<!-- grade_id: cell-13b1cef5f7407fdf, points: 50, grade: true, solution: true -->
YOUR ANSWER HERE

### Code tests

❓Use the next cells to include tests of your code. Call the functions that you have developed with known inputs and compare its return values with the expected results.



```python

```

```python

```