# DSP Lab 2022-23


This lab assignment requires preparation before the lab session.

Paragraphs that start with ❓explain what you will need to do in the following cell.

Before you turn this notebook in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).


---

## Proj3 - Pitch, beat and velocity

In this third part of the project, we intend to extract the remaining features necessary for the synthesis of the melody. Knowing the location of the onset, it is necessary to identify the note. This can be achieved by determining the fundamental frequency in a stable region of the signal following the onset. The synthesis also requires knowing how long the note should last. In music, the beat is used as the basic unit of time, with the duration of the note expressed in multiples or submultiples of its value. It is thus possible to change the rhythm of the music by modifying the beat value. Finally, it is necessary to estimate the intensity of each note. In the case of an instrument with a keyboard, the intensity depends on the speed with which the musician presses the keys. Modern synthesizers use a communication protocol called MIDI that encodes this speed in a 128-level scale. Dannenberg (2006) concludes that most manufacturers make an approximately linear relationship between the square of the speed and the peak RMS of the signal produced. As with the detection of onsets, the automatic detection of pitch, beat, and velocity must be tested on at least 3 different recordings, and the results compared with those obtained manually.

### Group and Student ID

❓Initialize the variable `group_id` with the number that Fenix assigned to your group and `ist_id1` and `ist_id2` with your student numbers.

<!-- grade_id: cell-2acee589c7979f8e, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()

print(f"Group number: {group_id}")
print(f"Student number: {ist_id1}")
print(f"Student number: {ist_id2}")
```

<!-- grade_id: cell-af624e6642d3e7fd, points: 10, grade: true, solution: false -->
```python
assert isinstance(group_id, int) and isinstance(ist_id1, int) and isinstance(ist_id2, int)
assert (group_id > 0) and (group_id < 35)
assert (ist_id1 > 60000) and (ist_id1 < 120000) and (ist_id2 > 60000) and (ist_id2 < 120000)
```

### Instructions and Description

❓Explain how to use your code and describe how it works.

<!-- grade_id: cell-8cfedd49f305e6c7, points: 50, grade: true, solution: true -->
YOUR ANSWER HERE

### Code

❓Place you code in the next cells. You can include additional code and text cells.



```python

```

```python

```

```python

```

### Assessment

❓Present the perfomance of each detector on the test signals and comment the achieved results.

<!-- grade_id: cell-13b1cef5f7407fdf, points: 50, grade: true, solution: true -->
YOUR ANSWER HERE

### Code tests

❓Use the next cells to include tests of your code. Call the functions that you have developed with known inputs and compare its return values with the expected results.



```python

```

```python

```